# Assignment 2: Do you see what I see?

The version of _WordyPy_ that we worked on in the last module was pretty basic, and with our newfound knowledge of image
manipulation it seems reasonable for us to increase the visual appeal here a bit. In this assignment I've modified the
`GameEngine` so that it returns an image to your bot -- your job is to be able to process that image and keep your bot
functioning!

## Requirements

The architecture for the software has changed a bit, here's a breakdown:

### `GameEngine`

I have changed the `GameEngine` class to now return a `PIL.Image` object when the `play` method is called. While it
doesn't sound like much of a change, it means that you need to use image manipulation to determine whether a letter was
in the correct place, in the word but not in the correct place, or not in the word at all.

### `DisplaySpecification`

While I demonstrate below how this looks using the Michigan blue and maize color scheme, I thought it would be nice to
allow future users of _WordPy_ to customize the look and feel, so your code needs to be able to get this new
`DisplaySpecification` and use it to solve the game. In this assignment the `GameEngine` will use the
`DisplaySpecification` defaults, so you can access the member attributes of the class directly.

### `Letter`

This class is unchanged, but you still need to provide the code for it! I recommend copying you solution for assignment
1 here.

### `Bot`

The biggest change to your `Bot` is that the `record_guess_results()` method no longer takes a list of letters, but an
image. To deal with this image you need to implement two helper methods which are private (remember: if a function or
method starts with an underscore in python the intention is for it to be private to the class). Our autograder will be
testing each of these functions to ensure that they work properly, so there's a brief discussion of them here.

#### Writing the method `Bot._tuple_to_str(self, pixels: tuple) -> str`

When you use the python imaging library to explore pixel values you generally get them back as a tuple, with each value
being an integer representation of a single color band. The `GameEngine` uses web hex color codes instead, so you'll
need to translate the tuple representation into a hex string equivalent. **This should be one of the first functions you
write.**

For example, the tuple values of `(186, 12, 47, 255)` in an RGBA image would convert to the hex string value of
`#BA0C2F`. Three things to note about your implementation of this function:

1. It must return hex values in upper case only
2. It must start with the hash mark (#) to indicate they are web hex colors
3. It must only consider the first three values in the tuple, and ignore any other values

#### Writing the method `Bot._process_image(self, guess: str, guess_image: PIL.Image) -> list[Letter]`

The `Bot._process_image` method takes two parameters, a `PIL.Image` which is the output from the game engine, and a
guess, as a string. This method needs to be wired up in your `record_guess_results()` method so that when the game
engine passes a new image you first decode it to a list of letters then record the results.

## Architecture

The new architecture diagram for WordyPy is the following (bold indicates new elements from assignment 1):
![Architecture Diagram of WordyPy for Assignment 2](py3_assignment2.drawio.png)

## Example

### Exploring the new `GameEngine`

It is helpful to know how to use the new `GameEngine` for development. In this assignment you can create a new
`GameEngine` for testing with your target word directly, then you can call its' `evaluate_guess()` method and pass in a
`Guess` for a work which you have chosen. This allows you to see the result in the Jupyter notebook directly using the
`display()` function. Here's a worked example:

```
# create new GameEngine with a target word of "field" and then
# show what it returns when we guess "lives"
guess = Guess("lives")
image = GameEngine("field").evaluate_guess(guess)
print(image)
display(image)
```

Example output: ![Example Output of GameEngine.evaluate_guess()](output1.png)

Note that to run this you _must_ first copy over your working `Letter` implementation from the first assignment, and you
_must_ run the cell below which has the new `GameEngine` implementation.

You will notice that when we print the PIL.Image object we get some information about its parameters (e.g. the image
mode and size). When we call the `display()` the object we see the word has one letter in the correct position (the "I",
shown on a dark blue background) and two letters which are in the word but in the incorrect position (the "L" and the
"E", shown on a maize colored background).

The `_tuple_to_str` private method of your `Bot` is one of the first ones you should write! Here is some code you can
use to test this function:

```python
# the pixel values as an RGBA series
pixels=(0,43,12,255)
print(Bot._tuple_to_str(pixels))
```

Here is an example of the output running on my solution for the assignment:
![Example Output of Bot._tuple_to_str() function](output2.png)


<!-- grade_id: cell-14420ae7de41900e, points: 0, grade: false, solution: true -->
```python
# Create the implementation of the Letter class here


# YOUR CODE HERE
class Letter:
    def __init__(self, letter: str, index: int = 0, in_word: bool = False, in_correct_place: bool = False) -> None:
        self.letter = letter.upper()
        self.in_correct_place = in_correct_place
        self.in_word = in_word
        self.index = index

    def is_in_correct_place(self) -> bool:
        return self.in_correct_place

    def is_in_word(self) -> bool:
        return self.in_word

    def __str__(self) -> str:
        return "Letter: {} idx({}) iCP({}) iW({})".format(self.letter, self.index, self.in_correct_place, self.in_word)
```

<!-- grade_id: cell-d87ca6087d3ca0b8, points: 1, grade: true, solution: false -->
```python
# This cell has the tests your Letter class

# Check if the Letter class exists
assert "Letter" in dir(), "The Letter class does not exist, did you define it?"

# Check to see if the Letter class can be created
l: Letter
try:
    l = Letter("s")
except:
    assert (
        False
    ), "Unable to create a Letter object with Letter('s'), did you correctly define the Letter class?"

# Check to see if the Letter class has the in_correct_place attribute
assert hasattr(
    l, "in_correct_place"
), "The letter object created has no in_correct_place attribute, but this should be False by default. Did you create this attribute?"

```

```python
class DisplaySpecification:
    block_width: int = 80
    block_height: int = 80
    correct_location_color: str = "#00274C"
    incorrect_location_color: str = "#FFCB05"
    incorrect_color: str = "#D3D3D3"
    space_between_letters: int = 5
    word_color: str = "#FFFFFF"
```

<!-- grade_id: cell-f52849cc3c29c39a, points: 0, grade: false, solution: true -->
```python
# Create the implementation of the Bot class here


# YOUR CODE HERE
import random
import PIL.Image

class Bot:

    def __init__(self, word_list_file: str, display_spec: DisplaySpecification = None) -> None:
        self.word_list_file: str = word_list_file
        with open(word_list_file, "r") as fh:
            raw: list = [line.strip().upper() for line in fh if line.strip()]
        self.word_list: list = list(dict.fromkeys(raw))
        self.display_spec: DisplaySpecification = display_spec

        print(f"\n*****\n{len(raw)} words loaded from file...")
        print(f"{len(self.word_list)} words without duplicates...")

        self.known_positions: list = [None, None, None, None, None]
        self.in_word_letters: set = set()
        self.excluded_letters: set = set()
        self.previous_guesses: list = []

    def _tuple_to_str(self, pixels: tuple) -> str:
        return '#' + (''.join(hex(pixel).replace('0x', '0').upper()[-2:] for pixel in pixels[:3]))

    def _process_image(self, guess: str, guess_image: PIL.Image.Image) -> list:
        spec = self.display_spec
        letters: list = []

        img_rgb = guess_image.convert("RGB")

        for i, ch in enumerate(guess.upper()):
            x: int = i * (spec.block_width + spec.space_between_letters) + spec.block_width // 2
            y: int = spec.block_height // 2

            pixel_tuple = img_rgb.getpixel((x, y))
            color_hex: str = self._tuple_to_str(pixel_tuple)

            in_correct_place: bool = (color_hex == spec.correct_location_color)
            in_word: bool = in_correct_place or (color_hex == spec.incorrect_location_color)

            letter = Letter(ch, index=i, in_word=in_word, in_correct_place=in_correct_place)
            letters.append(letter)

        return letters

    def _get_candidates(self) -> list:
        candidates = []
        for word in self.word_list:
            if word in self.previous_guesses:
                continue

            if any(ch in self.excluded_letters for ch in word):
                continue

            position_ok = all(
                self.known_positions[i] is None or word[i] == self.known_positions[i]
                for i in range(5)
            )
            if not position_ok:
                continue

            if any(ch not in word for ch in self.in_word_letters):
                continue

            candidates.append(word)

        return candidates

    def make_guess(self) -> str:
        candidates = self._get_candidates()

        if not candidates:
            candidates = [w for w in self.word_list if w not in self.previous_guesses]

        if not candidates:
            raise RuntimeError("No valid candidates remain – word list exhausted.")

        return random.choice(candidates)

    def record_guess_results(self, guess: str, guess_results: PIL.Image.Image) -> None:
        guess = guess.upper()
        self.previous_guesses.append(guess)

        letters = self._process_image(guess, guess_results)

        for idx, letter_obj in enumerate(letters):
            letter_obj.index = idx
            ch: str = letter_obj.letter.upper()

            if letter_obj.is_in_correct_place():
                self.known_positions[idx] = ch
                self.in_word_letters.add(ch)

            elif letter_obj.is_in_word():
                self.in_word_letters.add(ch)

            else:
                if ch not in self.known_positions and ch not in self.in_word_letters:
                    self.excluded_letters.add(ch)

    def __str__(self) -> str:
        return "Bot: File({}) TotalWords({}) KnownPos({}) Excluded({}) InWord({})".format(
            self.word_list_file,
            len(self.word_list),
            self.known_positions,
            self.excluded_letters,
            self.in_word_letters,
        )
```

```python
from PIL import Image, ImageFont, ImageDraw
import random

class GameEngine:

    def __init__(self, display_spec: DisplaySpecification = None) -> None:
        if display_spec == None:
            display_spec = DisplaySpecification()
        self.display_spec = display_spec

        self.err_input = False
        self.err_guess = False
        self.prev_guesses = []

    def play(
        self, bot: Bot, word_list_file: str = "words.txt", target_word: str = None
    ) -> Image:

        def format_results(results) -> str:
            response = ""
            for letter in results:
                if letter.is_in_correct_place():
                    response = response + letter.letter
                elif letter.is_in_word():
                    response = response + "*"
                else:
                    response = response + "?"
            return response

        word_list: list(str) = list(
            map(lambda x: x.strip().upper(), open(word_list_file, "r").readlines())
        )
        known_letters: list(str) = [None, None, None, None, None]
        unused_letters = set()

        if target_word is None:
            target_word = random.choice(word_list).upper()
        else:
            target_word = target_word.upper()
            if target_word not in word_list:
                print(f"Target word {target_word} must be from the word list")
                self.err_input = True
                return

        print(
            f"Playing a game of WordyPy using the word list file of {word_list_file}.\nThe target word for this round is {target_word}\n"
        )

        MAX_GUESSES = 6
        for i in range(1, MAX_GUESSES):
            guess: str = bot.make_guess()

            print(f"Evaluating bot guess of {guess}")

            if guess not in word_list:
                print(f"Guessed word {guess} must be from the word list")
                self.err_guess = True
            elif guess in self.prev_guesses:
                print(f"Guess word cannot be the same one as previously used!")
                self.err_guess = True

            if self.err_guess:
                return

            self.prev_guesses.append(guess)  # record the previous guess
            for j, letter in enumerate(guess):
                if letter in unused_letters:
                    print(
                        f"The bot's guess used {letter} which was previously identified as not used!"
                    )
                    self.err_guess = True
                if known_letters[j] is not None:
                    if letter != known_letters[j]:
                        print(
                            f"Previously identified {known_letters[j]} in the correct position is not used at position {j}!"
                        )
                        self.err_guess = True

                if self.err_guess:
                    return

            correct, results = self._set_feedback(guess, target_word)
            print(f"Was this guess correct? {correct}")
            img = self._format_results(results)
            print(f"Sending guess results to bot:\n")
            display(img)
            bot.record_guess_results(guess, img)

            if correct:
                print(f"Great job, you found the target word in {i} guesses!")
                return

        print(
            f"Thanks for playing! You didn't find the target word in the number of guesses allowed."
        )
        return

    def _set_feedback(self, guess: str, target_word: str) -> tuple[bool, list[Letter]]:
        correct: bool = True

        letters = []
        for j in range(len(guess)):
            letter = Letter(guess[j])

            if guess[j] == target_word[j]:
                letter.in_correct_place = True
            else:
                correct = False

            if guess[j] in target_word:
                letter.in_word = True

            letters.append(letter)

        return correct, letters

    def _render_letter(self, letter: Letter) -> Image:
        color: str = self.display_spec.incorrect_color
        if letter.is_in_correct_place():
            color = self.display_spec.correct_location_color
        elif letter.is_in_word():
            color = self.display_spec.incorrect_location_color

        block = Image.new(
            "RGB",
            (self.display_spec.block_width, self.display_spec.block_height),
            color=color,
        )
        draw = ImageDraw.Draw(block)

        X: int = self.display_spec.block_width // 2
        Y: int = self.display_spec.block_height // 2

        FONT_SIZE: int = 50
        font = ImageFont.truetype("assets/roboto_font/Roboto-Bold.ttf", FONT_SIZE)

        draw.text((X, Y), letter.letter, size=FONT_SIZE, anchor="mm", font=font)

        return block

    def _format_results(self, letters: list[Letter]) -> Image:
        WORD_WIDTH: int = (len(letters) * self.display_spec.block_width) + (
            len(letters) - 1
        ) * self.display_spec.space_between_letters
        WORD_HEIGHT: int = self.display_spec.block_height

        word = Image.new(
            "RGB", (WORD_WIDTH, WORD_HEIGHT), color=self.display_spec.word_color
        )
        curr_loc = 0
        for letter in letters:
            rendered_letter: Image = self._render_letter(letter)
            word.paste(rendered_letter, (curr_loc, 0))
            curr_loc += (
                self.display_spec.block_width + self.display_spec.space_between_letters
            )

        return word
```

<!-- grade_id: cell-02cee23cfeacbc7f, points: 1, grade: true, solution: false -->
```python
# Tests for Bot class.

# Check if the Bot class exists
assert "Bot" in dir(), "The Bot class does not exist, did you define it?"

```

```python
if __name__ == "__main__":
    # Chris's favorite words
    favorite_words = ["doggy", "drive", "daddy", "field", "state"]

    # Write this to a temporary file
    words_file = "words.txt"
    with open(words_file, "w") as file:
        file.writelines("\n".join(favorite_words))

    # Create a new GameEngine with the default DisplaySpecification
    ge = GameEngine()

    # Initialize the student Bot using the display specification from the game engine object
    bot = Bot(words_file, ge.display_spec)

    # Play a game with the Bot
    ge.play(bot, word_list_file=words_file)
```

<!-- grade_id: cell-10fc831637a3d6ac, points: 1, grade: true, solution: false -->
```python

```